**Wichtige Hinweise**

- Aktivieren Sie die entsprechende Conda-Umgebung bevor Sie beginnen.

- Geben Sie als `NAME` ihr HdM-Kürzel an.

- Ändern Sie **nicht** den Namen der Datei und löschen Sie keine Zellen.

- Bearbeiten Sie alle Zellen mit dem Hinweis <font color='green'> \# YOUR CODE HERE </font> 

- Die Funktion **NotImplementedError()** soll die Abgabe von leeren Zellen verhindern. Löschen Sie die Funktion, sobald Sie in einer dieser Zellen arbeiten. 

- Stellen Sie sicher, dass alles wie erwartet läuft, bevor Sie die Prüfung abgeben: Starten Sie den Kernel neu und führen Sie alle Zellen aus: wählen Sie "Restart" und dann "Run All" 

```python
NAME = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

---

# Introduction to linear regression

## Setup

```python
import pandas as pd
import altair as alt
import warnings

warnings.simplefilter(action='ignore', category=FutureWarning)
alt.data_transformers.disable_max_rows()
```

## Data

We create our own data:

### Create data

```python
df = pd.DataFrame(
    {'sales': [2500, 4500, 6500, 8500, 10500, 12500, 14500, 16500, 18500, 20500],
      'ads'  : [1000, 2000, 3000, 4000, 5000, 6000, 7000, 8000, 9000, 10000]}
)
```

### Data structure

```python
df
```

*Do you recognize a relationship between the two variables?*

```python
df.info()
```

## Analysis

Show the relationship between the variables:

```python
chart = alt.Chart(df).mark_point().encode(
   x=alt.X('ads', axis=alt.Axis(title='Ads (in $)')),
   y=alt.Y('sales', axis=alt.Axis(title="Sales (in units)")),
   tooltip=['ads', 'sales']
).interactive()

chart
```

## Model

Let's take a closer look at ad spendings of 2000. What is the value of sales you would expect?

```python
callout = alt.Chart(df.iloc[1:2]).mark_point(
    color='red', size=300, tooltip="Ads = 2000, Sales = 4500"
).encode(
    x=alt.X('ads', axis=alt.Axis(title='Ads (in $)')),
    y=alt.Y('sales', axis=alt.Axis(title="Sales (in units)"))
)

chart + callout
```

### Prediction

What is your sales prediction for tv ad spendings of 2000?

- calculate a prediction for sales using ad spendings of 2000
- in your code, use the variables `number_0` and `number_1` to obtain your prediction
- save the result as `sales_prediction`

Hint:

---

```python

number_0 = ___
number_1 = ___
ad_spendings = ___

sales_prediction = number_0 + number_1 * ad_spendings

```

---

<!-- grade_id: model, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c-model, points: 2, grade: true, solution: false -->
```python
# check your code
assert 4000 <= sales_prediction <= 5000
```

Next, use your solution to make your calculations within pandas for every value of ad spendings (`ads`) and save the result in your dataframe (as `sales_prediction`):

Hint:

---

```python
df['___'] = number_0 + number_1 * df['___'] 
```

---

- name the new column `sales_prediction`


<!-- grade_id: pred, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c-pred, points: 2, grade: true, solution: false -->
```python
# Check your code
assert 2000 <= df.iloc[0, 2] <= 3000
```

```python
df.head()
```

Visualize predictions as a line

```python
line = alt.Chart(df).mark_line().encode(
         alt.X('ads', axis=alt.Axis(title='Ads (in $)')),
         alt.Y('sales_prediction', axis=alt.Axis(title="Sales (in units)")),
         color=alt.value("#0001F5"))

chart + line
```

Next, we want to see where our blue line crosses the y axis. In other words, we need to include a value of `ads = 0` to our dataframe. 

Hint:

---

```python

df_new = pd.DataFrame({
    'ads': ___ , 
    'sales': ___, 
    'sales_prediction': ___ + ___ * ___
    }, ___)

___ = ___.___([___ , ___], ignore_index = ___)

```

---



First, we create a new Dataframe called `df_new` with only one row: 

- the value for `ads` is `0` 
- we don't know the real value of `sales`: therefore, we write `None` (Pandas will transform this to `NaN`, which means "Not a Number") in this cell 
- we use our formula from above to obtain the value of `sales_prediction` (use the value of `0` for ads)
- we need to manually include an index. Simply use `index=[0]`  

Next, we need to append `df_new` to the end of or DataFrame object `df`:

- use `pd.concat()` to combine `df` and `df_new` and save the result in `df` 
- we clear the existing index and reset it in the result by setting the `ignore_index` option to `True`.


<!-- grade_id: concat, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c-concat, points: 2, grade: true, solution: false -->
```python
# check your code
assert len(df) == 11
assert df.loc[10, 'ads'] == 0
```

```python
df
```

```python
chart_new = alt.Chart(df).mark_point().encode(
   x=alt.X('ads', axis=alt.Axis(title='Ads (in $)')),
   y=alt.Y('sales', axis=alt.Axis(title="Sales (in units)"))
)

line_new = alt.Chart(df).mark_line().encode(
         alt.X('ads', axis=alt.Axis(title='Ads (in $)')),
         alt.Y('sales_prediction', axis=alt.Axis(title="Sales (in units)")),
         color=alt.value("#0001F5"))

callout_new = alt.Chart(df.loc[10:'sales_prediction']).mark_point(
    color='red', 
    size=300, 
    tooltip="Ads = 0, Sales prediction = 500"
).encode(
    x='ads',
    y='sales_prediction'
)

chart_new + line_new + callout_new
```