**Wichtige Hinweise**

- Aktivieren Sie die entsprechende Conda-Umgebung bevor Sie beginnen.

- Geben Sie als `NAME` ihr HdM-Kürzel an.

- Ändern Sie **nicht** den Namen der Datei und löschen Sie keine Zellen.

- Bearbeiten Sie alle Zellen mit dem Hinweis <font color='green'> \# YOUR CODE HERE </font> 

- Die Funktion **NotImplementedError()** soll die Abgabe von leeren Zellen verhindern. Löschen Sie die Funktion, sobald Sie in einer dieser Zellen arbeiten. 

- Stellen Sie sicher, dass alles wie erwartet läuft, bevor Sie die Prüfung abgeben: Starten Sie den Kernel neu und führen Sie alle Zellen aus: wählen Sie "Restart" und dann "Run All" 

```python
NAME = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

---

# Find the best model

## Setup

```python
import pandas as pd
import altair as alt
import warnings

warnings.simplefilter(action='ignore', category=FutureWarning)
alt.data_transformers.disable_max_rows()
```

## Data

```python
# let's change some values in ads
df = pd.DataFrame(
    {'sales': [2500, 4500, 6500, 8500, 10500, 12500, 14500, 16500, 18500, 20500],
      'ads'  : [900, 1400, 3600, 3800, 6200, 5200, 6800, 8300, 9800, 10100]}
)
```

## Analysis

```python
alt.Chart(df).mark_point().encode(
    x='ads',
    y='sales',
    tooltip= ['ads', 'sales']
).interactive()
```

*Is there a "perfect" relationship between the two variables (like in our last example)?*

## Model

Remember our simple model from before. Instead of using `number_0` and `number_1`, we use the following expressions:

- `number_0` = `intercept`

- `number_1` = `slope` (also called coefficient)


Use the following values to make a prediction:

- intercept = 500
- slope = 2

Hint:

---


```python
intercept = ___ 
slope = ___

df['___'] = ___ + ___ * df['___'] 
```

---

- Make a prediction for `sales` by using the variable `ads`
- Save the new column as `sales_prediction`

<!-- grade_id: sales-prediction, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
df.head()
```

<!-- grade_id: c-sales-prediction, points: 4, grade: true, solution: false -->
```python
# Check your code
assert df.loc[0, 'sales_prediction'] == 2300 
```

```python
chart = alt.Chart(df).mark_point().encode(
    x='ads',
    y='sales'
)

line = alt.Chart(df).mark_line().encode(
         alt.X('ads', axis=alt.Axis(title='Ads (in $)')),
         alt.Y('sales_prediction', axis=alt.Axis(title="Sales (in units)")),
         color=alt.value("#0001F5"))

chart + line
```

*Can you guess why the two numbers are called intercept and slope?* Write down your answer below.