**Wichtige Hinweise**

- Aktivieren Sie die entsprechende Conda-Umgebung bevor Sie beginnen.

- Geben Sie als `NAME` ihr HdM-Kürzel an.

- Ändern Sie **nicht** den Namen der Datei und löschen Sie keine Zellen.

- Bearbeiten Sie alle Zellen mit dem Hinweis <font color='green'> \# YOUR CODE HERE </font> 

- Die Funktion **NotImplementedError()** soll die Abgabe von leeren Zellen verhindern. Löschen Sie die Funktion, sobald Sie in einer dieser Zellen arbeiten. 

- Stellen Sie sicher, dass alles wie erwartet läuft, bevor Sie die Prüfung abgeben: Starten Sie den Kernel neu und führen Sie alle Zellen aus: wählen Sie "Restart" und dann "Run All" 

```python
NAME = ""
```

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

---

# Fitting a line


## Setup

<!-- grade_id: setup, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Data

### Import data

Import https://raw.githubusercontent.com/kirenz/datasets/master/possum.csv and save as df

<!-- grade_id: df, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Data structure

```python

```

```python

```

### Variable lists

Prepara data for scikit-learn model:

use `head_l` (head length) as label (`y`) and total_l (total length) as predictor (`X`)

<!-- grade_id: var, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c-var, points: 2, grade: true, solution: false -->
```python
## Check your code
assert X.iloc[0,0] == 89
```

## Model

### Select model

<!-- grade_id: reg, points: 0, grade: false, solution: true -->
```python
# Choose the linear regression model
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c-reg, points: 2, grade: true, solution: false -->
```python
### Check your code
assert reg.get_params() == {'copy_X': True, 'fit_intercept': True, 'n_jobs': None, 'positive': False}
```

### Fit model

<!-- grade_id: reg-fit, points: 0, grade: false, solution: true -->
```python
# Fit the model to the data
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c-reg-fit, points: 2, grade: true, solution: false -->
```python
### Check your code
assert reg.positive == False
```

### Coefficients

Save the intercept with the name `intercept`

<!-- grade_id: int, points: 0, grade: false, solution: true -->
```python
# Intercept
# YOUR CODE HERE
raise NotImplementedError()
```

Save the coefficient with the name `slope`

<!-- grade_id: slope, points: 0, grade: false, solution: true -->
```python
# Slope

# YOUR CODE HERE
raise NotImplementedError()
```

### Make predictions

<!-- grade_id: pred, points: 0, grade: false, solution: true -->
```python
# Make predictions on the data (y_pred)
# YOUR CODE HERE
raise NotImplementedError()
```

```python
y_pred
```

### Evaluation

#### Mean squared error

Save the mean squarer error as `mse`

<!-- grade_id: mse, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Root mean squared error

Save the root mean squarer error as `rmse`

<!-- grade_id: rmse, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```